<a target="_blank" rel="noopener noreferrer" href="https://colab.research.google.com/github/epacuit/introduction-machine-learning/blob/main/tutorials/tutorial3.ipynb">![Open In Colab](https://colab.research.google.com/assets/colab-badge.svg)</a>


Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Andy Cheng"
COLLABORATORS = ""
```

---

# Tutorial 3: Brief Introduction to Pandas

This tutorial will provide a very brief introduction to the Pandas library. Pandas is a powerful data manipulation library for Python.

For a more in-depth introduction to Pandas, read the [Pandas Documentation](https://pandas.pydata.org/docs/).



#### Import the Pandas library and read a dataset

The first step is to read a dataset into a Pandas DataFrame. A DataFrame is a two-dimensional, size-mutable, and heterogeneous tabular data structure with labeled axes (rows and columns).

```python
import pandas as pd
```

```python
url = 'https://raw.githubusercontent.com/epacuit/introduction-machine-learning/refs/heads/main/tutorials/comedy_comparisons_metadata.csv'

df = pd.read_csv(url)
```

```python
type(df) # the type of the object is a DataFrame
```

Output:
```
pandas.core.frame.DataFrame
```

```python
df.head() # shows the first 5 rows of the DataFrame
```

Output:
```
video_id                                              title  view_count  \
0  DE1-cD3pTkA                      Walker Texas Ranger clip LIVE       224.0   
1  XZqSz_X-j8Y                Egg mcmuffin of a crappy commercial      1919.0   
2  vzpD6OogahQ           Potter Puppet Pals: School Is For Losers  14958231.0   
3  _OpzEHBDwQE  Very Funny Commercial  Think Beyond Your Ear (...     61696.0   
4  yzGWOpop6i8                               Happy Valentines Day     15852.0   

   like_count  comment_count duration  
0         1.0            0.0    PT32S  
1         7.0            3.0    PT44S  
2     87806.0        11321.0    PT33S  
3       179.0            2.0    PT38S  
4       315.0           77.0  PT1M14S
```

```python
df.columns # lists the columns of the DataFrame
```

Output:
```
Index(['video_id', 'title', 'view_count', 'like_count', 'comment_count',
       'duration'],
      dtype='object')
```

```python
df.shape # shows the shape of the DataFrame (rows, columns)
```

Output:
```
(11541, 6)
```

```python
len(df) # shows the number of rows in the DataFrame
```

Output:
```
11541
```

#### Creating a DataFrame



```python
df_1 = pd.DataFrame({
    'A': [1, 2, 3],
    'B': [4, 5, 6],
    'C': ['x', 'y', 'z']})

df_1
```

Output:
```
A  B  C
0  1  4  x
1  2  5  y
2  3  6  z
```

```python
df_2 = pd.DataFrame(
    [[1, 4, 'x'], [2, 5, 'y'], [3, 6, 'z']],
    columns=['A', 'B', 'C'])
df_2
```

Output:
```
A  B  C
0  1  4  x
1  2  5  y
2  3  6  z
```

#### Selecting Columns

The next step is to filter the DataFrame to select specific rows or columns.

```python
df['title'] # A column can be accessed by using the column name as a key
```

Output:
```
0                            Walker Texas Ranger clip LIVE
1                      Egg mcmuffin of a crappy commercial
2                 Potter Puppet Pals: School Is For Losers
3        Very Funny Commercial  Think Beyond Your Ear (...
4                                     Happy Valentines Day
                               ...                        
11536              Dylan hit by a ball [Slow motion remix]
11537            Trick or Treat! Funny Halloween Dog Video
11538                                   Framed Thief Prank
11539           Harry Redknapp gets hit by Fernando Torres
11540                    ~600 Subs Special~ MMV- Enchanted
Name: title, Length: 11541, dtype: object
```

```python
type(df['title']) # the type of a column is a Series
```

Output:
```
pandas.core.series.Series
```

```python
titles = df['title'] # the column can be stored in a variable

titles.head() # shows the first 5 rows of the Series
```

Output:
```
0                        Walker Texas Ranger clip LIVE
1                  Egg mcmuffin of a crappy commercial
2             Potter Puppet Pals: School Is For Losers
3    Very Funny Commercial  Think Beyond Your Ear (...
4                                 Happy Valentines Day
Name: title, dtype: object
```

```python
print(f"The 11th title is {titles[10]}") # a value in a Series can be accessed by using the index as a key

print(f'The 11th title is {df["title"][10]}') # you can also access a value in a Series by using the DataFrame and column name

print("\nThe first 5 values of the title: ", titles[0:5]) # A Series object can be sliced like a list

print("\nThe last 5 values of the title: ", titles[-5:]) # A Series object can be sliced like a list

```

Output:
```
The 11th title is PCD Flying Fun Cam
The 11th title is PCD Flying Fun Cam

The first 5 values of the title:  0                        Walker Texas Ranger clip LIVE
1                  Egg mcmuffin of a crappy commercial
2             Potter Puppet Pals: School Is For Losers
3    Very Funny Commercial  Think Beyond Your Ear (...
4                                 Happy Valentines Day
Name: title, dtype: object

The last 5 values of the title:  11536       Dylan hit by a ball [Slow motion remix]
11537     Trick or Treat! Funny Halloween Dog Video
11538                            Framed Thief Prank
11539    Harry Redknapp gets hit by Fernando Torres
11540             ~600 Subs Special~ MMV- Enchanted
Name: title, dtype: object
```

```python
print("The unique values of the view_counts column are: ", df['view_count'].unique()) # shows the unique values of a column

print("The number of values of the view_counts column are: ", df['view_count'].count()) # shows the number of values of a column

print("The number of unique values of the view_counts column are: ", df['view_count'].nunique()) # shows the number of unique values of a column

print(" ", len(list(set(df['view_count'])))) # Another way to find the number of unique elements in a column
```

Output:
```
The unique values of the view_counts column are:  [2.2400000e+02 1.9190000e+03 1.4958231e+07 ... 1.3728710e+06 1.7410000e+03
 5.9710000e+03]
The number of values of the view_counts column are:  11532
The number of unique values of the view_counts column are:  9310
  9319
```

Typically, it is faster to use the Pandas's built-in functions.

```python
%%timeit

df['view_count'].nunique()
```

Output:
```
468 s  7.43 s per loop (mean  std. dev. of 7 runs, 1000 loops each)
```

```python
%%timeit

len(list(set(df['view_count'])))
```

Output:
```
2.13 ms  431 s per loop (mean  std. dev. of 7 runs, 1000 loops each)
```

#### Filtering the DataFrame

```python
df[df['view_count'] > 1000000] # shows the rows that have a view_count greater than 1,000,000
```

Output:
```
video_id                                              title  \
2      vzpD6OogahQ           Potter Puppet Pals: School Is For Losers   
18     rQjr5xjugUM  Ah meu Deus do cu, Quebraram(quebraru) a mesa...   
23     U2_CSYI8O0Y                              Kud - Koninginnemarkt   
25     TugslL45aXk                           Don't argue with an ibex   
39     TXtP8UfIfEA      JibJab  Valentine's Day eCards.flv_(360p).mp4   
...            ...                                                ...   
11506  a91eiu_eer4                        Optical Illusion Girlfriend   
11516  59xYAKjCnIo                                        SUPER TIMOR   
11530  q-RLqLx1iYI             Epic and Honest Mobile Home Commercial   
11531  Ej9wCko6dR4                   Bud Light Jackie Moon Commercial   
11538  dHuPgeClQC0                                 Framed Thief Prank   

       view_count  like_count  comment_count duration  
2      14958231.0     87806.0        11321.0    PT33S  
18      4364246.0     
```

```python
df['view_count'] > 1000000 # returns a boolean Series
```

Output:
```
0        False
1        False
2         True
3        False
4        False
         ...  
11536    False
11537    False
11538     True
11539    False
11540    False
Name: view_count, Length: 11541, dtype: bool
```

```python
df[df['comment_count'].isin([5, 10, 1000])] # shows the rows that have a comment count that is either 5, 10, or 1000 greater than 1,000,000
```

Output:
```
video_id                                     title  view_count  \
28     rBVoffKg6QA               Talking Cats (High Version)     42287.0   
54     jDD8yq18Gaw  19:57 From Euston - WHAT REALLY HAPPENED     19296.0   
109    XqkjPnrqbSc    dahk fokaha tarjama ( chehal fi Sa3a )     85511.0   
149    v-kHJjqLya4                      Funny Dog Commercial     65501.0   
201    -3g_ENHcvGY                          Grave yard Smash      1369.0   
...            ...                                       ...         ...   
11174  AxQK9pp8nVc         louie c.k. balls in your pus-hole      2765.0   
11313  -UzGgGJkIFU                         JB SKETCH BAD BOY     14321.0   
11392  DiA8VItyPIk                 Tryin first gear wheelies      1181.0   
11396  NhHRfS5bWug                          Natanek  2013 AD       614.0   
11429  4DQxWeTuG7I     Family Guy Peter and the Osama Clawed      3195.0   

       like_count  comment_count duration  
28           91.0           10.0    PT56S  
54       
```

```python
df[(df['view_count'] > 1000000) & (df['comment_count'] > 1000)] # shows the rows that have a view_count greater than 1,000,000 and a comment_count greater than 1000
```

One issue to be aware of is that when you filter a DataFrame, the index of the original DataFrame is preserved.  


```python
df[df["view_count"] % 2 == 0] # get all rows where the view_count is even
```

Output:
```
video_id                                              title  \
0      DE1-cD3pTkA                      Walker Texas Ranger clip LIVE   
3      _OpzEHBDwQE  Very Funny Commercial  Think Beyond Your Ear (...   
4      yzGWOpop6i8                               Happy Valentines Day   
6      DOD38Y72zyk                              Minecraft: TNT Rocket   
8      c1NaGio2Geo                                 Caller X episode 3   
...            ...                                                ...   
11524  gJyAigxDWBI  Knowwheretorun1984 set ablaze by Jordan Maxwel...   
11529  Sn00l_ttSi4                    "Strangers on the Street" - QCC   
11532  XcVw-qHdufE  Trailer: Being Human - Series 4 Episode 2 - BB...   
11535  K-w_Ie-iBgY                                Gears of War 3 Rage   
11536  HKMgwmqPOxs            Dylan hit by a ball [Slow motion remix]   

       view_count  like_count  comment_count duration  
0           224.0         1.0            0.0    PT32S  
3         61696.0    
```

```python
# get the 1 element of the view_count column
print("The 2nd element of the view_count column is ", df.loc[1, 'view_count'])
df[df["view_count"] % 2 == 0][1] # produces an error since the item with index 1 does not have an even view_count
```

```python
df[df["view_count"] % 2 == 0].values[1] # use the values attribute to get the values in the filtered dataframe.
```

#### Statistics about the DataFrame

```python
df["view_count"].sum() # shows the mean of the view_count column
```

```python
sum(df["view_count"]) # outputs nan because the column has missing values (listed at NaN)
```

```python
sum(df["view_count"].dropna()) # outputs the sum of the view_count column without the missing values
```

Again, it is faster to use the Pandas's built-in functions.

```python
%%timeit

df["view_count"].sum()
```

```python
%%timeit

sum(df["view_count"].dropna())
```

```python
df["like_count"].max() # shows the max of the view_count column
```

```python
df.describe() # shows the summary statistics of the DataFrame
```

#### Combining DataFrames

```python
df_3 = pd.DataFrame({
    'A': [10, 11, 12],
    'B': [13, 14, 15],
    'C': ['xx', 'yy', 'zz']})

df_3
```

Output:
```
A   B   C
0  10  13  xx
1  11  14  yy
2  12  15  zz
```

```python
combined_df = pd.concat([df_1, df_2, df_3], axis=0) # concatenates the two DataFrames along the rows

combined_df
```

Output:
```
A   B   C
0   1   4   x
1   2   5   y
2   3   6   z
0   1   4   x
1   2   5   y
2   3   6   z
0  10  13  xx
1  11  14  yy
2  12  15  zz
```

```python
combined_df = pd.concat([df_1, df_2, df_3], axis=1) # concatenates the two DataFrames along the columns

combined_df
```

### Task

<!-- grade_id: cell-aefcc9d3857dc581, points: 0, grade: false, solution: true -->
```python
def compare_views(df, video_id1, video_id2):
    """Return True if the view count of video_id1 is greater than the view count of video_id2; otherwise, return False. If video_id1 or video_id2 is not in the DataFrame, raise a ValueError that outputs the string 'video_id is not in the DataFrame', where video_id is not present in the frame."""

    # df_videoid1 = df.loc[df['video_id']] == video_id1
    # df_videoid2 = df.loc[df['video_id']] == video_id2
    # print(df['video_id']).values

    try:
      video_1 = int(df.loc[df['video_id'] == video_id1, 'view_count'].values)
    except:
      raise ValueError(f'{video_id1} is not in the DataFrame')

    try:
      video_2 = int(df.loc[df['video_id'] == video_id2, 'view_count'].values)
    except:
      raise ValueError(f'{video_id2} is not in the DataFrame')

    if video_1 > video_2:
        return True
    elif video_1 < video_2:
        return False
    elif video_1 == video_2:
        return False
    else:
        raise ValueError(f'{video_id1} is not in the DataFrame')
```

<!-- grade_id: cell-a0439c6c3ac29144, points: 7, grade: true, solution: false -->
```python
assert compare_views(df, 'vzpD6OogahQ', 'yzGWOpop6i8') == True
assert compare_views(df, 'yzGWOpop6i8', 'vzpD6OogahQ') == False
assert compare_views(df, 'DE1-cD3pTkA', 'XZqSz_X-j8Y') == False
assert compare_views(df, 'XZqSz_X-j8Y', 'DE1-cD3pTkA') == True
assert compare_views(df, 'yzGWOpop6i8', 'yzGWOpop6i8') == False
try:
    compare_views(df, 'vzpD6OogahQ', 'not_in_df1')
except ValueError as e:
    assert str(e) == 'not_in_df1 is not in the DataFrame'
try:
    compare_views(df, 'not_in_df2', 'vzpD6OogahQ')
except ValueError as e:
    assert str(e) == 'not_in_df2 is not in the DataFrame'
```

Output:
```
<ipython-input-87-b5bab163dedb>:9: DeprecationWarning: Conversion of an array with ndim > 0 to a scalar is deprecated, and will error in future. Ensure you extract a single element from your array before performing this operation. (Deprecated NumPy 1.25.)
  video_1 = int(df.loc[df['video_id'] == video_id1, 'view_count'].values)
<ipython-input-87-b5bab163dedb>:14: DeprecationWarning: Conversion of an array with ndim > 0 to a scalar is deprecated, and will error in future. Ensure you extract a single element from your array before performing this operation. (Deprecated NumPy 1.25.)
  video_2 = int(df.loc[df['video_id'] == video_id2, 'view_count'].values)
```

```python

```